# Python для сбора и анализа данных

*Разработчик формата сдачи заданий и фабрики автоматических тестов: Щуров И.В.<br> 
Автор задач: Тамбовцева А.А.*

Задача считается решенной, если код успешно проходит автоматические тесты и не содержит грубых нарушений («подгонка» под тесты, копирование кода другого студента). «Подгонка» под тесты – написание кода не для общего случая, а для соответствия конкретным значениям в ячейке с тестами. Пример: пишем не универсальный код для проверки четности числа, а код, который правильно работает только на числах 2, 5, 8, потому что именно эти числа используются в `test_data`. 

### Формат сдачи

Для сдачи домашнего задания необходимо скачать текущий ipynb-файл, вписать в него свои решения, сохранить изменения и загрузить обновленный файл на Dropbox по предложенной на странице курса ссылке. Иметь аккаунт Dropbox для этого не требуется, достаточно пройти по ссылке и выбрать файл на компьютере.

Получить ipynb-файл можно, выбрав в Jupyter пункт меню *File → Download as... → Notebook (.ipynb)*.

### Тесты для проверки кода

На первом этапе проверки ваши решения в домашних заданиях и лабораторных работах будут проверяться с помощью автоматических тестов – ячеек с кодом, которые проверяют соответствие результатов исполнения ваших программ ожидаемым результатам в условиях задач. Эти тесты можно использовать для проверки кода перед сдачей задания. 

Итак, для предварительной проверки задания нужно сделать следующее:

1. Скачать данный ipynb-файл на свой компьютер, открыть его в Jupyter Notebook.
2. Активировать тесты (см. ниже).
3. Вставить решение каждой задачи в ячейку для кода, следующую за его условием, там, где написан комментарий `# YOUR CODE HERE`. Отступ вашего кода должен составлять 4 пробела. Ничего не менять вокруг!
4. Запустить ячейку, в которую вы вставили код с решением. Ввести какие-то входные данные, проверить визуально правильность вывода.
5. Запустить следующую ячейку (в ней содержится тест). Если запуск ячейки с тестом не приводит к появлению ошибки (`assertion`), значит, всё в порядке, задача решена. Если приводит к появлению ошибки, значит, тест не пройден и нужно искать ошибку.

Внимание! Если в какой-то момент забыть ввести входные данные и перейти на следующую ячейку, есть риск, что Notebook перестанет откликаться. В этом случае надо перезапустить ядро: *Kernel → Restart*. При этом потеряются все значения переменных, но сам код останется.

### Активировать тесты

Запустите следующую ячейку, чтобы иметь возможность запускать тесты.

```python
# Фабрика тестов для проверки программ, принимающих данные через input()

from collections import deque

class Tester(object):
    def __init__(self, inp):
        self.outputs = []
        self.inputs = deque(inp)
    def print(self, *args, sep = " ", end = "\n"):
        text = sep.join(map(str, args)) + end
        newlines = text.splitlines(keepends=True)
        if self.outputs and self.outputs[-1] and self.outputs[-1][-1] != "\n" and newlines:
            self.outputs[-1] += newlines[0]
            self.outputs.extend(newlines[1:])
        else:
            self.outputs.extend(newlines)
            
    def input(self, *args):
        assert self.inputs, "Вы пытаетесь считать больше элементов, чем предусмотрено условием"
        return self.inputs.popleft()
    def __enter__(self):
        global print
        global input
        print = self.print
        input = self.input
        return self.outputs
    def __exit__(self, *args):
        global print
        global input
        del print
        del input
```

### Задача 1. В синем море, в белой пене... (1 балл)

В Светлогорске, в музее Мирового океана есть прекрасные весы, которые позволяют узнать свой вес, измеренный в селёдках, в китах, в креветках и в других морских обитателях.

Напишите программу, которая запрашивает у пользователя с клавиатуры его вес в килограммах (целое или дробное число) и выводит на экран его вес в селёдках, округлённый до целого числа. Считайте, что средний вес селёдки равен 350 граммам.

Уточнение: считайте, что пользователь вводит допустимые значения – целые и дробные числа больше 0. Задача не предполагает проверки ввода пользователя, но если интересно, можете над этим подумать.

**Пример работы программы**

Ввод:

    56

Вывод:

    Ваш вес в селёдках: 160.

<!-- grade_id: ans02, points: 0, grade: false, solution: true -->
```python
def more():
    
    # YOUR CODE HERE

more()
```

<!-- grade_id: test02, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("56", "Ваш вес в селёдках: 160."),
    ("35", "Ваш вес в селёдках: 100."),
    ("64", "Ваш вес в селёдках: 183."),
    ("82", "Ваш вес в селёдках: 234."),
    ("40.5", "Ваш вес в селёдках: 116."),
    ("47.25", "Ваш вес в селёдках: 135.")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        more()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 2. На дальней станции сойду... (1 балл)

Дан список станций желтой ветки московского метро (упорядочены как на карте):

    line = ["Третьяковская", "Марксистская", "Площадь Ильича", "Авиамоторная", "Шоссе Энтузиастов", "Перово", "Новогиреево", "Новокосино"]

Давайте считать, что движение происходит от Третьяковской до Новокосино.

Напишите программу, которая запрашивает у пользователя название текущей станции метро (без лишнего пробела на конце) и выводит на экран сообщение вида

    Следующая станция: [станция].

Считайте, что пользователь знает, какая станция конечная, и не будет запрашивать станцию, следующую после Новокосино.

**Пример**

Входные данные:

    Площадь Ильича

Выходные данные:

    Следующая станция: Авиамоторная.

<!-- grade_id: ans01, points: 0, grade: false, solution: true -->
```python
def metro():
    line = ["Третьяковская", "Марксистская", "Площадь Ильича", "Авиамоторная", "Шоссе Энтузиастов",
    "Перово", "Новогиреево", "Новокосино"]
    
    # YOUR CODE HERE

metro()
```

<!-- grade_id: test01, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("Третьяковская", "Следующая станция: Марксистская."),
    ("Марксистская", "Следующая станция: Площадь Ильича."),
    ("Площадь Ильича", "Следующая станция: Авиамоторная."),
    ("Авиамоторная", "Следующая станция: Шоссе Энтузиастов."),
    ("Шоссе Энтузиастов","Следующая станция: Перово."),
    ("Перово", "Следующая станция: Новогиреево."),
    ("Новогиреево", "Следующая станция: Новокосино.")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        metro()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 3. Любопытство – не порок (1 балл)

Любопытный питон Петя ловит прохожих на улице и просит назвать их какую-нибудь цифру от 0 до 9. Все ответы он записывает в виде последовательности из цифр, а затем считает сумму этих цифр. Итак, напишите программу, которая:

* сначала запрашивает число прохожих с клавиатуры (число $n$);
* затем $n$ раз просит ввести какую-нибудь цифру;
* выводит на экран полученную последовательность цифр и их сумму через пробел.

Считайте, что пользователь всегда вводит с клавиатуры корректные значения, целые положительные числа для количества прохожих и цифры от 0 до 9.

Пользоваться готовой функцией `sum()` для вычисления суммы элементов списка нельзя.

**Примеры работы программы**

Ввод:

    4
    0
    9
    8
    2

Вывод:

    0982 19
    
Ввод:

    3
    1
    1
    5

Вывод:

    115 7
    
Ввод:

    7
    9
    9
    8
    1
    2
    3
    0

Вывод:

    9981230 32

<!-- grade_id: ans02, points: 0, grade: false, solution: true -->
```python
def curious():
    
    # YOUR CODE HERE

curious()
```

<!-- grade_id: test02, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("4 5 6 0 1", "5601 12"),
    ("7 9 9 8 1 2 3 0", "9981230 32"),
    ("2 0 1", "01 1"),
    ("3 1 1 5", "115 7"),
    ("10 3 4 7 7 3 1 0 2 3 5", "3477310235 35"),
    ("1 0", "0 0")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        curious()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, были введены числа " + inp
```

### Задача 4. Чем больше, тем веселее (1 балл)

В массиве `valid` сохранено число действительных бюллетеней на 5 избирательных участках, а в массиве `invalid` – число недействительных бюллетеней на тех же участках. Также известно общее число избирателей на каждом участке, значения сохранены в массиве `total`.

```python
import numpy as np

valid = np.array([450, 800, 630, 740, 950])
invalid = np.array([10, 200, 100, 50, 120])
total = np.array([900, 1000, 1200, 890, 2300])
```

Используя имеющиеся массивы, вычислите явку (в процентах) на каждом избирательном участке, сохраните результаты в массив `turnout` и выведите на экран с новой строки его элементы, округлённые до второго знака после запятой.

*Напоминание.* Явка на каждом участке определяется как сумма действительных и недействительных бюллетеней, делённая на общее число зарегистированных избирателей (не забудьте перевести полученную долю в проценты).

**Примеры работы программы**

На выходе:

    51.11
    96.02
    60.83
    88.76
    46.52

<!-- grade_id: ans04, points: 0, grade: false, solution: true -->
```python
def elections():
    
    valid = np.array([450, 787, 630, 740, 950])
    invalid = np.array([10, 202, 100, 50, 120])
    total = np.array([900, 1030, 1200, 890, 2300])
    
    # YOUR CODE HERE
    
elections()
```

<!-- grade_id: test04, points: 1, grade: true, solution: false -->
```python
test_data = [("", ["51.11\n", "96.02\n", "60.83\n",
                   "88.76\n", "46.52\n"])]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        elections()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверное количество строк в выводе"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), \
                "Неверное количество элементов в строке " + l_out
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Ошибка {} != {}".format(l_test, 
                                                                   l_out)
```

### Задача 5. Пора вспомнить про друзей в Австралии (1 балл)

У питона Пети в записной книжке хранятся полные имена друзей и знакомых из разных стран. Записная книжка представляет собой словарь:

    pythons = {"королевский питон Рональд" : "Уганда", 
      "карликовый питон Антониа" : "Ангола", 
      "сетчатый питон Асиф" : "Бангладеш", 
      "зеленый питон Оливер" : "Австралия", 
      "тиморский питон Николау" : "Восточный Тимор", 
      "аметистовый питон Харрисон" : "Австралия", 
      "тигровый питон Мэйли" : "Китай",
      "оливковый питон Руби" : "Австралия", 
      "белогубый питон Арис" : "Индонезия", 
      "эфиопский питон Оника" : "Мозамбик"}

Питон Петя хочет съездить в Австралию, и поэтому хочет вывести на экран имена тех друзей и знакомых, которые живут в Австралии (чтобы потом выбрать, у кого остановиться). Напишите программу, которая выводит на экран полные имена питонов из Австралии.

**Ожидаемый вывод на экран:**

    зеленый питон Оливер
    аметистовый питон Харрисон
    оливковый питон Руби

```python
def petya_au():
    
    pythons = {"королевский питон Рональд" : "Уганда", 
    "карликовый питон Антониа" : "Ангола", 
    "сетчатый питон Асиф" : "Бангладеш", 
    "зеленый питон Оливер" : "Австралия", 
    "тиморский питон Николау" : "Восточный Тимор", 
    "аметистовый питон Харрисон" : "Австралия", 
    "тигровый питон Мэйли" : "Китай",
    "оливковый питон Руби" : "Австралия", 
    "белогубый питон Арис" : "Индонезия", 
    "эфиопский питон Оника" : "Мозамбик"}
    
    # YOUR CODE HERE
    

petya_au()
```

```python
test_data = [
    ("", ["зеленый питон Оливер\n", "аметистовый питон Харрисон\n", "оливковый питон Руби\n"])
]

for inp, out in test_data:
    with Tester([inp]) as t:
        petya_au()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверное количество строк в выводе"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), \
                "Неверное количество элементов в строке " + l_out
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Ошибка {} != {}".format(l_test, 
                                                                   l_out)
```

### Задача 6. Мультпарад (1 балл)

Реализуйте программу, находящую мультфильмы по имени героя:

* Дан словарь `data`, в котором ключами являются имена героев, а значениями — списки названий мультфильмов, где такие герои встречаются. Гарантируется, что названия мультфильмов не повторяются, все имена и названия записаны в нижнем регистре.
* В переменную `role` считывается имя героя. Этот код уже написан. Гарантируется, что такой герой есть в словаре.
* Программа выводит названия тех, мультфильмов, где такой герой есть, но его имя не фигурирует в названии. Если подходящих названий нет — программа ничего не выводит.


**Примеры работы программы:**

Ввод:

    ежик

Результат:

    осенние корабли
    трям! здравствуйте!
    
Ввод:

    ослик

Результат:

    заветная мечта
    винни-пух и день забот
    
Ввод:

    енот

Результат:

    [ничего не выводится]

```python
def mp2():

    data = {"ежик" : ["ежик в тумане", "осенние корабли", 
                      "трям! здравствуйте!", "ежик и девочка"], 
    "баба-яга" : ["баба-яга против!", "ивашка из дворца пионеров", 
                  "василиса прекрасная"], 
    "ослик" : ["как ослик грустью заболел", "козлик и ослик", 
               "заветная мечта", "винни-пух и день забот"],
    "волк" : ["серый волк энд красная шапочка", "самый маленький гном", 
             "волк и теленок", "ну, погоди"],
    "енот" : ["крошка енот"]}

    role = input()

    # YOUR CODE HERE
    
mp2()
```

```python
test_data = [
    ("ежик", ["осенние корабли", "трям! здравствуйте!"]),
    ("баба-яга", ["ивашка из дворца пионеров", "василиса прекрасная"]),
    ("ослик", ["заветная мечта", "винни-пух и день забот"]),
    ("волк", ["самый маленький гном", "ну, погоди"]),
    ("енот", [])
]

for inp, out in test_data:
    with Tester(inp.split(",")) as t:
        mp2()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверный ответ"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), "Неверный ответ"
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Неверный ответ" 
```

### Задача 7. Города, города... (1 балл)

Дан словарь с записями, ключами которых являются числовые id, а значениями – города.

```python
cities = {100: 'Москва', 101: 'Санкт-Петербург', 
          103: 'Екатеринбург', 104: 'Пермь', 
          107: 'Красноярск', 111 : 'Киров'}
```

Напишите программу, которая просит пользователя ввести с клавиатуры сначала ключ, потом значение, а затем:

* если в словаре `cities` есть запись с таким ключом, на экран выводится сообщение "Запись с таким ключом существует. Введите другой ключ."

* если в словаре `cities` нет записи с таким ключом, в этот словарь добавляется запись с введенным ключом и значением + обновленный словарь `cities` выводится на экран.

**Пример:**

*Входные данные:*

    Введите ключ: 101
    Введите значение: Мурманск

*Выходные данные:*

    Запись с таким ключом существует. Введите другой ключ.

*Входные данные:*

    Введите ключ: 117
    Введите значение: Омск

*Выходные данные (порядок записей может быть другим):*

    {100: 'Москва', 101: 'Санкт-Петербург', 
    103: 'Екатеринбург', 104: 'Пермь', 
    107: 'Красноярск', 111 : 'Киров', 117 : 'Омск'}

Обратите внимание, id должны быть целочисленными!

```python
def goroda():
    
    cities = {100: 'Москва', 101: 'Санкт-Петербург', 
          103: 'Екатеринбург', 104: 'Пермь', 
          107: 'Красноярск', 111 : 'Киров'}
    
    # YOUR CODE HERE
    
goroda()
```

```python
test_data = [
    ("100 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("101 a", "Запись с таким ключом существует. Введите другой ключ."), 
    ("103 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("104 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("107 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("111 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("115 Омск", "{100: 'Москва', 101: 'Санкт-Петербург', 103: 'Екатеринбург', 104: 'Пермь', 107: 'Красноярск', 111: 'Киров', 115: 'Омск'}"),
    ("1 Казань", "{100: 'Москва', 101: 'Санкт-Петербург', 103: 'Екатеринбург', 104: 'Пермь', 107: 'Красноярск', 111: 'Киров', 1: 'Казань'}")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        goroda()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 8. Снова метро (2 балла)

Дан словарь, содержащий информацию по станциям метро в Новосибирске (станции в списке указаны последовательно, с севера на юг в случае Ленинской ветки, и с запада на восток в случае Дзержинской ветки).

```python
nmetro = {"Ленинская" : ["Заельцовская", "Гагаринская", "Красный проспект", 
                         "Площадь Ленина", "Октябрьская",
                           "Речной Вокзал", "Студенческая", 
                         "Площадь Маркса"], 
          "Дзержинская" : ["Площадь Гарина-Михайловского",
                           "Сибирская", "Маршала Покрышкина",
                           "Березовая Роща", "Золотая Нива"]}
```

Напишите программу, которая запрашивает у пользователя сначала название текущей станции без пробела на конце, затем направление движения (`'1'` – с севера на юг или с запада на восток, `'2'` – с юга на север или с востока на запад) и выводит на экран название следующей станции.

Считайте, что пользоватль знает, какие станции являются конечными, и не будет запрашивать станцию, следующую после конечной.

**Пример:**

*Входные данные:*

    Введите название текущей станции: Гагаринская
    Выберите направление движения (1 - с севера на юг или с запада на восток, 2 - с юга на север или с востока на запад): 2

*Выходные данные:*

    Следующая станция: Заельцовская


*Входные данные:*

    Введите название текущей станции: Сибирская
    Выберите направление движения (1 - с севера на юг или с запада на восток, 2 - с юга на север или с востока на запад): 1

*Выходные данные:*

    Следующая станция: Маршала Покрышкина

```python
def metro():
    
    nmetro = {"Ленинская" : ["Заельцовская", "Гагаринская", "Красный проспект", 
                         "Площадь Ленина", "Октябрьская",
                           "Речной Вокзал", "Студенческая", 
                         "Площадь Маркса"], 
          "Дзержинская" : ["Площадь Гарина-Михайловского",
                           "Сибирская", "Маршала Покрышкина",
                           "Березовая Роща", "Золотая Нива"]}
    
    # YOUR CODE HERE

metro()
```

```python
test_data = [
    ("Гагаринская;2", "Следующая станция: Заельцовская"),
    ("Гагаринская;1", "Следующая станция: Красный проспект"), 
    ("Сибирская;1", "Следующая станция: Маршала Покрышкина"),
    ("Сибирская;2", "Следующая станция: Площадь Гарина-Михайловского"),
    ("Речной Вокзал;1", "Следующая станция: Студенческая"),
    ("Речной Вокзал;2", "Следующая станция: Октябрьская"),
]

for inp, out in test_data:
    with Tester(inp.split(";")) as t:
        metro()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 9. Добрые дела (2 балла)

Питон Петя решил сделать хорошее дело – раздать те книги, которые у него хранятся в нескольких экземлярах. 

Напишите функцию `to_give()`, которая:

* принимает на вход список книг `books` (список строк);
* возвращает отсортированный по алфавиту список книг, которые нужно отдать, то есть те книги, которые в списке встречаются более одного раза. 

ВАЖНО! Функция не должна ничего выводить на экран.

**Примеры работы функции**

Ввод:

    ['Хобби доктора Травена', 'Хобби доктора Травена', 'Змеелов', 'Маленький принц', 
    'Питоны о Питонах. Автобиография', 'Python для детей', 'Питон для детей', 'Кролики и удавы', 
    'Маленький принц', 'Удавы и питоны. Уход и содержание', 'Питон для детей']

Результат:

    ['Маленький принц', 'Питон для детей', 'Хобби доктора Травена']

Ввод:

    ['Змеелов', 'Маленький принц', 'Хобби доктора Травена']

Результат:

    []    

<!-- grade_id: p3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: p3-tests, points: 2, grade: true, solution: false -->
```python
assert to_give(['Хобби доктора Травена', 
                'Хобби доктора Травена', 
                'Змеелов', 'Маленький принц', 
                'Питоны о Питонах. Автобиография', 
                'Python для детей', 'Питон для детей', 
                'Кролики и удавы', 'Маленький принц', 
                'Удавы и питоны. Уход и содержание', 
                'Питон для детей']) == ['Маленький принц', 'Питон для детей', 
                                        'Хобби доктора Травена'], 'Неверно'

assert to_give(['Змеелов', 'Маленький принц', 'Хобби доктора Травена']) == [], 'Неверно'
assert to_give(['C', 'C', 'D', 'E', 'A', 'A', 'A', 'A']) == ['A', 'C'], 'Неверно'
```

### Задача 10. Учат в школе (3 балла)

Напишите функцию `eq_solve()`, которая принимает на вход строку с квадратным уравнением (без пробелов между знаками и со всеми коэффициентами, записанными явно, включая 1 и -1) и решает квадратное уравнение:

* если дискриминант отрицательный, возвращает пустой список;
* если дискриминант равен нулю, возвращает список из одного решения;
* если дискриминант больше нуля, возвращает список из двух решений, упорядоченных по возрастанию.

Считайте, что квадратное уравнение уже приведено к стандартному виду:

$$
ax^2+bx+c=0
$$

При решении нельзя использовать библиотечные функции (кроме функции для извлечения квадратного корня) и регулярные выражения.

**Примеры работы функции**

Применение:

    eq_solve('1x^2+6x+9=0')


Возвращаемый результат:

    [-3.0]

Применение:

    eq_solve('-2x^2-6x+9=0')

Возвращаемый результат:

    [-4.098076211353316, 1.098076211353316]

Применение:

    eq_solve('3x^2-1x+10=0')


Возвращаемый результат:

    []

<!-- grade_id: p6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: p6-tests, points: 3, grade: true, solution: false -->
```python
from math import isclose

assert isclose(eq_solve('1x^2+6x+9=0')[0], -3.0), "Неверное решение при D = 0"
assert isclose(eq_solve('-2x^2-6x+9=0')[0], -4.098076211353316), "Неверное решение при D > 0"
assert isclose(eq_solve('-2x^2-6x+9=0')[1], 1.098076211353316), "Неверное решение при D > 0"
assert len(eq_solve('3x^2-1x+10=0')) == 0, "Неверное решение при D < 0"
assert isclose(eq_solve('17x^2-5x-7=0')[0], -0.5112655672235116), "Неверное решение при D > 0"
assert isclose(eq_solve('17x^2-5x-7=0')[1], 0.8053832142823351), "Неверное решение при D > 0"
assert isclose(eq_solve('-10x^2+25x-4=0')[0], 0.1718070673576088), "Неверное решение при D > 0"
assert isclose(eq_solve('-10x^2+25x-4=0')[1], 2.3281929326423914), "Неверное решение при D > 0"
assert len(eq_solve('9x^2+4x+4=0')) == 0, "Неверное решение при D < 0"
```